[1512744860] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;1;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1512744920] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;2;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1512744980] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;3;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1512745040] SERVICE ALERT: localhost;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1512745357] SERVICE ALERT: Local Agent;CPU;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_nrpe, ...) failed. errno is 2: No such file or directory
[1512745403] SERVICE ALERT: Local Agent;DNS Service Uptime;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_nrpe, ...) failed. errno is 2: No such file or directory
[1512745417] SERVICE ALERT: Local Agent;CPU;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_nrpe, ...) failed. errno is 2: No such file or directory
[1512745449] SERVICE ALERT: Local Agent;HARD DISK Utilization;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_nrpe, ...) failed. errno is 2: No such file or directory
[1512745463] SERVICE ALERT: Local Agent;DNS Service Uptime;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_nrpe, ...) failed. errno is 2: No such file or directory
[1512745477] SERVICE ALERT: Local Agent;CPU;CRITICAL;SOFT;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_nrpe, ...) failed. errno is 2: No such file or directory
[1512745495] SERVICE ALERT: Local Agent;Network Traffic;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_nrpe, ...) failed. errno is 2: No such file or directory
[1512745509] SERVICE ALERT: Local Agent;HARD DISK Utilization;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_nrpe, ...) failed. errno is 2: No such file or directory
[1512745523] SERVICE ALERT: Local Agent;DNS Service Uptime;CRITICAL;SOFT;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_nrpe, ...) failed. errno is 2: No such file or directory
[1512745537] SERVICE ALERT: Local Agent;CPU;CRITICAL;HARD;4;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_nrpe, ...) failed. errno is 2: No such file or directory
[1512745555] SERVICE ALERT: Local Agent;Network Traffic;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_nrpe, ...) failed. errno is 2: No such file or directory
[1512745569] SERVICE ALERT: Local Agent;HARD DISK Utilization;CRITICAL;SOFT;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_nrpe, ...) failed. errno is 2: No such file or directory
[1512745583] SERVICE ALERT: Local Agent;DNS Service Uptime;CRITICAL;HARD;4;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_nrpe, ...) failed. errno is 2: No such file or directory
[1512745615] SERVICE ALERT: Local Agent;Network Traffic;CRITICAL;SOFT;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_nrpe, ...) failed. errno is 2: No such file or directory
[1512745629] SERVICE ALERT: Local Agent;HARD DISK Utilization;CRITICAL;HARD;4;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_nrpe, ...) failed. errno is 2: No such file or directory
[1512745675] SERVICE ALERT: Local Agent;Network Traffic;CRITICAL;HARD;4;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_nrpe, ...) failed. errno is 2: No such file or directory
[1512746342] HOST ALERT: Apache Web Server;DOWN;SOFT;1;check_ping: Invalid hostname/address - address
[1512746402] HOST ALERT: Apache Web Server;DOWN;SOFT;2;check_ping: Invalid hostname/address - address
[1512746402] SERVICE ALERT: Apache Web Server;Apache Web Server Status;UNKNOWN;SOFT;1;check_http: Invalid hostname/address - address
[1512746462] SERVICE ALERT: Apache Web Server;PING;UNKNOWN;SOFT;1;check_ping: Invalid hostname/address - address
[1512746462] SERVICE ALERT: Apache Web Server;Apache Web Server Status;UNKNOWN;SOFT;2;check_http: Invalid hostname/address - address
[1512746462] HOST ALERT: Apache Web Server;DOWN;SOFT;3;check_ping: Invalid hostname/address - address
[1512746522] HOST ALERT: Apache Web Server;DOWN;SOFT;4;check_ping: Invalid hostname/address - address
[1512746522] SERVICE ALERT: Apache Web Server;PING;UNKNOWN;SOFT;2;check_ping: Invalid hostname/address - address
[1512746522] SERVICE ALERT: Apache Web Server;Apache Web Server Status;UNKNOWN;SOFT;3;check_http: Invalid hostname/address - address
[1512746582] SERVICE ALERT: Apache Web Server;Apache Web Server Status;UNKNOWN;HARD;4;check_http: Invalid hostname/address - address
[1512746582] SERVICE ALERT: Apache Web Server;PING;UNKNOWN;SOFT;3;check_ping: Invalid hostname/address - address
[1512746582] HOST ALERT: Apache Web Server;DOWN;SOFT;5;check_ping: Invalid hostname/address - address
[1512746642] SERVICE ALERT: Apache Web Server;PING;UNKNOWN;HARD;4;check_ping: Invalid hostname/address - address
[1512746642] HOST ALERT: Apache Web Server;DOWN;SOFT;6;check_ping: Invalid hostname/address - address
[1512746702] HOST ALERT: Apache Web Server;DOWN;SOFT;7;check_ping: Invalid hostname/address - address
[1512746762] HOST ALERT: Apache Web Server;DOWN;SOFT;8;check_ping: Invalid hostname/address - address
[1512746822] HOST ALERT: Apache Web Server;DOWN;SOFT;9;check_ping: Invalid hostname/address - address
[1512746882] HOST ALERT: Apache Web Server;DOWN;HARD;10;check_ping: Invalid hostname/address - address
[1512747343] HOST ALERT: Apache Web Server;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1512747375] SERVICE ALERT: Apache Web Server;Apache Web Server Status;OK;HARD;4;HTTP OK: HTTP/1.1 200 OK - 11595 bytes in 0.003 second response time
[1512747389] SERVICE ALERT: Apache Web Server;PING;OK;HARD;4;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1512747646] HOST ALERT: MySQL;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1512747673] HOST ALERT: Apache Web Server;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1512747685] SERVICE ALERT: Apache Web Server;Apache Web Server Status;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1512747695] SERVICE ALERT: Apache Web Server;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1512747715] HOST ALERT: Apache Web Server;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1512747735] SERVICE ALERT: MySQL;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1512747736] HOST ALERT: MySQL;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1512747743] SERVICE ALERT: MySQL;MySQL Status;CRITICAL;SOFT;1;(Service check timed out after 61.01 seconds)
[1512747745] SERVICE ALERT: Apache Web Server;Apache Web Server Status;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1512747755] SERVICE ALERT: Apache Web Server;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1512747773] HOST ALERT: MySQL;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1512747775] HOST ALERT: Apache Web Server;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1512747795] SERVICE ALERT: MySQL;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1512747804] SERVICE ALERT: MySQL;MySQL Status;CRITICAL;SOFT;2;(Service check timed out after 61.01 seconds)
[1512747805] SERVICE ALERT: Apache Web Server;Apache Web Server Status;CRITICAL;SOFT;3;CRITICAL - Socket timeout
[1512747815] SERVICE ALERT: Apache Web Server;PING;CRITICAL;SOFT;3;PING CRITICAL - Packet loss = 100%
[1512747825] HOST ALERT: MySQL;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1512747835] HOST ALERT: Apache Web Server;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1512747855] SERVICE ALERT: MySQL;PING;CRITICAL;SOFT;3;PING CRITICAL - Packet loss = 100%
[1512747865] SERVICE ALERT: Apache Web Server;Apache Web Server Status;CRITICAL;HARD;4;CRITICAL - Socket timeout
[1512747865] SERVICE ALERT: MySQL;MySQL Status;CRITICAL;SOFT;3;(Service check timed out after 61.02 seconds)
[1512747875] SERVICE ALERT: Apache Web Server;PING;CRITICAL;HARD;4;PING CRITICAL - Packet loss = 100%
[1512747885] HOST ALERT: MySQL;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1512747895] HOST ALERT: Apache Web Server;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1512747915] SERVICE ALERT: MySQL;PING;CRITICAL;HARD;4;PING CRITICAL - Packet loss = 100%
[1512747926] SERVICE ALERT: MySQL;MySQL Status;CRITICAL;HARD;4;(Service check timed out after 61.01 seconds)
[1512747945] HOST ALERT: MySQL;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1512747985] HOST ALERT: Apache Web Server;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1512748035] HOST ALERT: MySQL;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1512748075] HOST ALERT: Apache Web Server;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1512748125] HOST ALERT: MySQL;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1512748165] HOST ALERT: Apache Web Server;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1512748195] HOST ALERT: Apache Web Server;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1512748215] HOST ALERT: MySQL;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1512748245] HOST ALERT: MySQL;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1512748285] HOST ALERT: Apache Web Server;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1512764441] SERVICE ALERT: Local Agent;HARD DISK Utilization;OK;HARD;4;DISK OK - free space: / 6334 MB (80% inode=92%):
[1512764486] SERVICE ALERT: Local Agent;CPU;OK;HARD;4;OK - user: , nice: , sys: , iowait: , irq: , softirq:  idle:
[1512764541] SERVICE ALERT: Local Agent;DNS Service Uptime;OK;HARD;4;/usr/sbin/named -f -u bind     01:24:23
[1512764553] SERVICE ALERT: Local Agent;Network Traffic;OK;HARD;4;RX Bytes: 17MB, TX Bytes:  2MB: RX Speed: 0Bps, TX Speed: 0Bps: OK bandwidth utilization